```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller3.ipynb")
```

<img src="Archivos/miad4.png" width=800x>

# Funciones

En este taller usarás funciones y funciones anónimas para resolver diversos problemas.

## Habilidades en práctica
Al desarrollar este taller podrás verificar tu progreso para:

**1.** Implementar funciones e integrarlas con condicionales y ciclos. <br>
**2.** Implementar funciones anónimas e integrarlas con condicionales.

## Instrucciones
En cada uno de los siguientes ejercicios deberás escribir el código solicitado estrictamente en las celdas indicadas para ello, teniendo en cuenta las siguientes recomendaciones:

* No crear, eliminar o modificar celdas de este Notebook (salvo lo que se te indique), pues puede verse afectado el proceso de calificación automática.
* La calificación se realiza de manera automática con datos diferentes a los proporcionados en este taller. Por consiguiente, tu código debe funcionar para diferentes instancias de cada uno de los ejercicios; una instancia hace referencia al valor de los parámetros.
* La calificación de cada ejercicio depende del valor que retorne la función especificada en su enunciado. Por lo tanto, aunque implementes funciones adicionales, es escencial que utilices los nombres propuestos en los enunciados de los ejercicios para implementar la función definitiva.

Esta es una actividad calificada y, por lo tanto, debe ser resuelta individualmente.

A continuación puedes ver un ejemplo de los ejercicios que encontrarás en este taller.

### Ejemplo

Implementa una función llamada `multiplicacion` que reciba dos números ingresados por parámetro y retorne la multiplicación de estos dos.

*Esta función debería retornar un número*

```python
# Implementa tu respuesta en esta celda

def multiplicacion(num1, num2):
    
    respuesta = num1 * num2
    
    return respuesta
```

### Ejercicio 1

A continuación, encuentras declarada una lista de números.

```python
# No modifiques esta celda

lista = [20, 12, 21.3, 34]
```

Implementa una función llamada `suma_lista` que reciba por parámetro una lista de números (por ejemplo, la variable `lista`) y retorne la suma de todos sus elementos.

*Esta función debe retornar un número*.

```python
def suma_lista(lista):
    # Escribe tu codigo aqui
    ...
```

```python
grader.check("q1")
```

### Ejercicio 2

A continuación, encuentras declarada una lista que contiene letras y números.

```python
lista = [23, "2", 46 , "l", 8.3, "m"]
```

Implementa una función llamada `primera_letra` que reciba una lista de números y letras por parámetro (por ejemplo, la variable `lista`), y retorne el primer elemento que sea una letra. La función debe retornar una cadena de texto vacía (`""`) si no encuentra una letra.

*Esta función debe retornar una letra*.

```python
def primera_letra(lista):
    # Escribe tu codigo aqui
    ...
```

```python
grader.check("q2")
```

### Ejercicio 3

A continuación, encuentras declarada una lista que contiene letras y números.

```python
lista = [4, 3, 8, 5, 6, 2, 7, 5]
```

Implementa una función llamada `separar_pares_impares` que reciba una lista de números enteros por parámetro (por ejemplo, la variable `lista`) y retorne una tupla que contenga dos listas, cada una ordenada ascendentemente:

* La primera lista debe contener los números pares.
* La segunda lista debe contener los números impares.

*Esta función debe retornar una tupla de listas*.

```python
def separar_pares_impares(lista):
    # Escribe tu codigo aqui
    ...
```

```python
grader.check("q3")
```

<!-- grade_id: ej3-test, points: 1, grade: true, solution: false -->
```python


# Base variables
lista = [4, 3, 8, 5, 6, 2, 7, 5]
lista_prueba = [3, 1, 2, 8, 6, 6, 7, 10, 9, 8]

# Caso 1: no existe la función.
try:
    separar_pares_impares
    assert type(separar_pares_impares) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada separar_pares_impares.")

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    separar_pares_impares(lista)
    separar_pares_impares(lista_prueba)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna una tupla de listas
assert type(separar_pares_impares(lista_prueba)) == tuple, f"Tu función debe retornar un valor de tipo {tuple.__name__}."
assert type(separar_pares_impares(lista_prueba)[0]) == list and type(separar_pares_impares(lista_prueba)[1]) == list, "La tupla que retorna tu función debe contener una lista en su primera y segunda posicion."

# Case 4: len(tupla) != 2
assert len(separar_pares_impares(lista_prueba)) == len(([],[])), "Tu función retorna una tupla de tamaño incorrecto."

# Caso 5: respuesta explicita
assert separar_pares_impares(lista_prueba) == ([2, 6, 6, 8, 8, 10], [1, 3, 7, 9]), "Tu respuesta es incorrecta para para una instancia diferente. Utiliza el parámetro."

# Caso 6: la lista de impares está ubicada antes que la lista de pares.
assert sum([i % 2 for i in separar_pares_impares(lista_prueba)[0]]) == 0, "La primera lista de la tupla debe contener únicamente pares."
assert sum([(i + 1) % 2 for i in separar_pares_impares(lista_prueba)[1]]) == 0, "La segunda lista de la tupla debe contener únicamente impares."

# Caso 7: la lista de pares y la lista de impares es cada una de longitud incorrecta
assert len(separar_pares_impares(lista_prueba)[0]) == len([2, 6, 6, 8, 8, 10]), "La primera lista de la tupla que retorna tu función no tiene la longitud correcta."
assert len(separar_pares_impares(lista_prueba)[1]) == len([1, 3, 7, 9]), "La segunda lista de la tupla que retorna tu función no tiene la longitud correcta."

# Caso 8: las listas son correctas, independiente de su orden
assert set(separar_pares_impares(lista_prueba)[0]) == {2, 6, 6, 8, 8, 10}, "La primera lista de la tupla que retorna tu función no contiene los elementos esperados, independiente de su orden."
assert set(separar_pares_impares(lista_prueba)[1]) == {1, 3, 7, 9}, "La segunda lista de la tupla que retorna tu función no contiene los elementos esperados, independiente de su orden."

# Caso 9: las listas no están ordenadas ascendentemente
assert separar_pares_impares(lista_prueba)[0] == [2, 6, 6, 8, 8, 10], "La primera lista de la tupla que retorna tu función no está ordenada ascendentemente."
assert separar_pares_impares(lista_prueba)[1] == [1, 3, 7, 9], "La segunda lista de la tupla que retorna tu función no está ordenada ascendentemente."

assert separar_pares_impares(lista)[0] == [2, 4, 6 ,8], "La primera lista de la tupla que retorna tu función no está ordenada ascendentemente."
assert separar_pares_impares(lista)[1] == [3, 5, 5, 7], "La segunda lista de la tupla que retorna tu función no está ordenada ascendentemente."

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

### Ejercicio 4

A continuación, encuentras declarados algunos números enteros.

```python
treinta = 30
cuarenta_y_cinco = 45
cincuenta = 50
sesenta_y_tres = 63
```

Implementa una función anónima que reciba un número entero por parámetro (por ejemplo, la variable `trinta`) y retorne `True` si el número es mayor a 50. Guarda la función (sin ejecutar) en una variable llamada `mayor_a_cincuenta`.

*Esta función debe retornar un valor booleano* (`True` o `False`).

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q4")
```

### Ejercicio 5

A continuación, encuentras declarada una lista de números enteros.

```python
lista = [3, 4, 2, 5, 5, 6, 3]
```

Implementa una función llamada `modulo_del_minimo` que reciba una lista de números enteros por parámetro (por ejemplo, la variable `lista`), encuentre el valor más pequeño que se repita en la lista (`min_val_rep`) y retorne una función anónima. Esta función anónima debe recibir un número entero mayor a cero (`entero`) por parámetro y retornar el residuo de la división entre `min_val_rep` y `entero` (recomendamos usar el operador de residuo `%`).

*Esta función debe retornar una función anónima*.

##### Ejemplo
```python
    anon = modulo_del_minimo([3, 4, 2, 5, 5, 6, 3])
    
    anon(2) == 1
```
Una vez implementada la función `modulo_del_minimo`, ejectar el código anterior en una celda de Jupyter debe generar el *output* `True`.

```python
def modulo_del_minimo(lista):
    # Escribe tu codigo aqui
    ...
```

```python
grader.check("q5")
```

## Créditos

__Autor(es)__: Camilo Falla Moreno, Alejandro Mantilla Redondo, Juan David Reyes, Diego Alejandro Cely Gómez
 
 __Fecha última actualización__: 06/08/2021

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```