```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller3.ipynb")
```

# Taller 3: Funciones

En este taller practicarás la implementación de funciones y funciones anónimas (lambda) en Python.

### Ejercicio 1

A continuación encuentras una lista de números. Implementa una función llamada `suma_lista` que reciba por parámetro una lista de números (por ejemplo, la variable `lista`) y retorne la suma de todos sus elementos.

*Esta función debe retornar un número.*

```python
# No modifiques esta celda
lista = [20, 12, 21.3, 34]
```

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q1")
```

### Ejercicio 2

A continuación encuentras una lista de números y letras. Implementa una función llamada `primera_letra` que reciba una lista de números y letras por parámetro (por ejemplo, la variable `lista`) y retorne el primer elemento que sea una letra. La función debe retornar una cadena de texto vacía (`""`) si no encuentra una letra.

*Esta función debe retornar una letra.*

```python
# No modifiques esta celda
lista = [23, "2", 46 , "l", 8.3, "m"]
```

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q2")
```

### Ejercicio 3

A continuación encuentras una lista de números enteros. Implementa una función llamada `separar_pares_impares` que reciba una lista de números enteros por parámetro (por ejemplo, la variable `lista`) y retorne una tupla que contenga dos listas, cada una ordenada ascendentemente:

* La primera lista debe contener los números pares.
* La segunda lista debe contener los números impares.

*Esta función debe retornar una tupla de listas.*

```python
# No modifiques esta celda
lista = [4, 3, 8, 5, 6, 2, 7, 5]
```

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q3")
```

### Ejercicio 4

A continuación encuentras unos números enteros. Implementa una función anónima que reciba un número entero por parámetro (por ejemplo, la variable `treinta`) y retorne `True` si el número es mayor a 50. Guarda la función (sin ejecutar) en una variable llamada `mayor_a_cincuenta`.

*Esta función debe retornar un valor booleano (`True` o `False`).*

```python
# No modifiques esta celda
treinta = 30
cuarenta_y_cinco = 45
cincuenta = 50
sesenta_y_tres = 63
```

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q4")
```

### Ejercicio 5

A continuación encuentras una lista de números enteros. Implementa una función llamada `modulo_del_minimo` que reciba una lista de números enteros por parámetro (por ejemplo, la variable `lista`), encuentre el valor más pequeño que se repita en la lista (`min_val_rep`) y retorne una función anónima. Esta función anónima debe recibir un número entero mayor a cero (`entero`) por parámetro y retornar el residuo de la división entre `min_val_rep` y `entero` (recomendamos usar el operador de residuo `%`).

*Esta función debe retornar una función anónima.*

##### Ejemplo
```python
anon = modulo_del_minimo([3, 4, 2, 5, 5, 6, 3])
anon(2) == 1
```

```python
# No modifiques esta celda
lista = [3, 4, 2, 5, 5, 6, 3]
```

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q5")
```

## Entrega

Asegurate de haber ejecutado todas las celdas de tu notebook en orden antes de ejecutar la celda de abajo. La celda generara un archivo ZIP para que lo entregues. **Guarda tu notebook antes de exportar.**

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert callable(suma_lista), 'suma_lista debe ser una funcion.'
>>> assert round(suma_lista([20, 12, 21.3, 34]), 2) == 87.3, 'Falla con la lista del enunciado.'
>>> assert suma_lista([1, 2, 3, 4, 5]) == 15, 'Falla con lista [1, 2, 3, 4, 5].'
```


Tests for q2 (points: 1)
```
>>> assert callable(primera_letra), 'primera_letra debe ser una funcion.'
>>> assert primera_letra([23, '2', 46, 'l', 8.3, 'm']) == 'l', 'Falla con la lista del enunciado.'
>>> assert primera_letra([1, 2, '3', 4.5]) == '', 'Falla cuando no hay letras.'
```


Tests for q3 (points: 1)
```
>>> assert callable(separar_pares_impares), 'separar_pares_impares debe ser una funcion.'
>>> assert separar_pares_impares([4, 3, 8, 5, 6, 2, 7, 5]) == ([2, 4, 6, 8], [3, 5, 5, 7]), 'Falla con la lista del enunciado.'
```


Tests for q4 (points: 1)
```
>>> assert callable(mayor_a_cincuenta), 'mayor_a_cincuenta debe ser callable.'
>>> assert mayor_a_cincuenta(63) is True, '63 debe dar True.'
>>> assert mayor_a_cincuenta(50) is False, '50 debe dar False (estrictamente mayor).'
>>> assert mayor_a_cincuenta(30) is False, '30 debe dar False.'
```


Tests for q5 (points: 1)
```
>>> assert callable(modulo_del_minimo), 'modulo_del_minimo debe ser una funcion.'
>>> fn = modulo_del_minimo([3, 4, 2, 5, 5, 6, 3])
>>> assert str(fn).find('lambda') > -1, 'modulo_del_minimo debe retornar una funcion lambda.'
>>> assert fn(2) == 1, 'En [3,4,2,5,5,6,3], el minimo repetido es 3. 3 % 2 debe ser 1.'
>>> assert fn(3) == 0, '3 % 3 debe ser 0.'
```
