```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller7.ipynb")
```

```python
import numpy as np
import pandas as pd
from datetime import datetime
```

```python
# No modifiques esta celda
orden_cliente = pd.DataFrame({'id_cliente': [1, 2, 3], 'monto': [100, 200, 300]})
cliente_categoria = pd.DataFrame({'id_cliente': [2, 3, 4], 'categoria': ['A', 'B', 'C']})
```

```python
def union_derecha(df1, df2):
    ...
...
```

```python
grader.check("q1")
```

```python
# No modifiques esta celda
cliente_categoria_orden = pd.DataFrame({'id': [1,2], 'val': [3,4]})

```

```python
def concatenar_interno(df1, df2):
    ...
...
```

```python
grader.check("q2")
```

```python
# No modifiques esta celda
ventas = pd.DataFrame({'A': [1, np.nan, 3]})

```

```python
def mediana_faltantes(df):
    ...
...
```

```python
grader.check("q3")
```

```python
# No modifiques esta celda
clientes = pd.DataFrame({'Sexo': ['Masculino', 'Femenino', 'Femenino'], 'Cliente Premium': ['Si', 'No', 'No']})

```

```python
def cliente_dummies(df):
    ...
...
```

```python
grader.check("q4")
```

```python
# No modifiques esta celda
fecha_string = '2020-10-December'

```

```python
def cadena_a_fecha(cadena):
    ...
...
```

```python
grader.check("q5")
```

```python
def union_izquierda():
    ...
...
```

```python
grader.check("q6")
```

```python
# No modifiques esta celda
df = pd.DataFrame({'no_orden': [1, np.nan, np.nan], 'idVendedor': [np.nan, 2, np.nan], 'otro': [1, 2, 3]})

```

```python
def ejercicio7(df):
    ...
...
```

```python
grader.check("q7")
```

```python
def ejercicio8():
    ...
...
```

```python
grader.check("q8")
```

## Entrega

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Asegurate de haber ejecutado todas las celdas...

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert union_derecha(orden_cliente, cliente_categoria).shape == (6, 3)
```


Tests for q2 (points: 1)
```
>>> assert concatenar_interno(cliente_categoria_orden, cliente_categoria).shape == (7, 2)
```


Tests for q3 (points: 1)
```
>>> assert mediana_faltantes(ventas).isnull().sum().sum() == 0
```


Tests for q4 (points: 1)
```
>>> assert list(cliente_dummies(clientes)['Sexo']) == [0, 1, 1]
```


Tests for q5 (points: 1)
```
>>> assert cadena_a_fecha(fecha_string) == datetime(2020, 12, 10, 0, 0)
```


Tests for q6 (points: 1)
```
>>> assert 'Fecha' in union_izquierda().columns
```


Tests for q7 (points: 1)
```
>>> assert ejercicio7(df) == ['no_orden', 'idVendedor']
```


Tests for q8 (points: 1)
```
>>> assert len(ejercicio8()) == 171
```
