```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller2.ipynb")
```

# Taller 2: Estructuras de control: condicionales y ciclos

Pregunta 1. Mayor de dos números. Variables: num_1=2, num_2=3. Encuentra el mayor.

```python
num_1 = 2
num_2 = 3
# No modifiques esta celda
```

```python
mayor = ...
...
```

```python
grader.check("q1")
```

Pregunta 2. Elemento en lista. Variables: elemento='e', lista=['b','c','a','d']. Agrega si no está.

```python
elemento = 'e'
lista = ['b','c','a','d']
# No modifiques esta celda
```

```python
# Escribe tu código aquí
...
...
```

```python
grader.check("q2")
```

Pregunta 3. División o residuo. Variables: num=7, divisor=3.

```python
num = 7
divisor = 3
# No modifiques esta celda
```

```python
division_o_residuo = ...
...
```

```python
grader.check("q3")
```

Pregunta 4. Corredor más rápido.

```python
vel_Camilo=10.5
vel_Esteban=9.3
vel_Alejandro=9.8
vel_David=10.2
# No modifiques esta celda
```

```python
mas_rapido = ...
...
```

```python
grader.check("q4")
```

Pregunta 5. Ocupación de un teatro.

```python
num_actual=20
cap_maxima=45
# No modifiques esta celda
```

```python
info = ...
...
```

```python
grader.check("q5")
```

Pregunta 6. Tipo de pizza.

```python
vegetarianos=['Tomate','Champiñones','Cebolla']
no_vegetarianos=['Pepperoni','Pollo','Jamón']
ingrediente_1='Champiñones'
ingrediente_2='Pollo'
ingrediente_3='Cebolla'
# No modifiques esta celda
```

```python
tipo = ...
...
```

```python
grader.check("q6")
```

Pregunta 7. Números pares.

```python
n=100
# No modifiques esta celda
```

```python
pares = ...
...
```

```python
grader.check("q7")
```

Pregunta 8. Cuenta regresiva.

```python
inicio=10
# No modifiques esta celda
```

```python
cuenta_regresiva = ...
...
```

```python
grader.check("q8")
```

Pregunta 9. Búsqueda de índice.

```python
numeros=[1,2,8,4,89,3,89]
a_buscar=89
# No modifiques esta celda
```

```python
indice = ...
...
```

```python
grader.check("q9")
```

Pregunta 10. Conteo de vocales.

```python
cadena='Me encanta programar en Python'
# No modifiques esta celda
```

```python
num_vocales = ...
...
```

```python
grader.check("q10")
```

Pregunta 11. Perros para evento.

```python
lb=2
ub=5
perros_refugio={'Max':1,'Bella':4,'Charlie':6,'Luna':3,'Rocky':5}
# No modifiques esta celda
```

```python
perros_evento = ...
...
```

```python
grader.check("q11")
```

Pregunta 12. Descifrar mensaje.

```python
cifrada='Aqxzxobjmp'
corr=3
alfabeto='abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'
# No modifiques esta celda
```

```python
descifrada = ...
...
```

```python
grader.check("q12")
```

Pregunta 13. Elección de regalo.

```python
catalogo={'Libro': 50000, 'Juego': 120000, 'Audífonos': 80000}
presupuesto=100000
# No modifiques esta celda
```

```python
regalo = ...
...
```

```python
grader.check("q13")
```

## Entrega

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Asegúrate de haber ejecutado todas las celdas...

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert mayor == 3
```


Tests for q10 (points: 1)
```
>>> assert num_vocales['a'] == 4
```


Tests for q11 (points: 1)
```
>>> assert set(perros_evento) == {'Bella', 'Luna', 'Rocky'}
```


Tests for q12 (points: 1)
```
>>> assert type(descifrada) == str
```


Tests for q13 (points: 1)
```
>>> assert type(regalo) == str
```


Tests for q2 (points: 1)
```
>>> assert 'e' in lista
>>> assert len(lista) == 5
```


Tests for q3 (points: 1)
```
>>> assert division_o_residuo == 1
```


Tests for q4 (points: 1)
```
>>> assert mas_rapido == 'Esteban'
```


Tests for q5 (points: 1)
```
>>> assert info == ['Bajo', 'Verde']
```


Tests for q6 (points: 1)
```
>>> assert tipo == 'Mixta'
```


Tests for q7 (points: 1)
```
>>> assert len(pares) == 50
>>> assert pares[-1] == 100
```


Tests for q8 (points: 1)
```
>>> assert cuenta_regresiva == [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]
```


Tests for q9 (points: 1)
```
>>> assert indice == 4
```
