```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller10.ipynb")
```

```python
import numpy as np
import pandas as pd
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LinearRegression
import pickle
```

## Pregunta 1

```python
archivo = pd.read_excel('Archivos/datosTaller.xlsx')
```

```python
def dividir_datos(df):
    ...
...
```

```python
grader.check("q1")
```

## Pregunta 2

```python
datos = (X, y)
```

```python
def crear_train_test(datos):
    ...
...
```

```python
grader.check("q2")
```

## Pregunta 3

```python
datos = (X_tr, X_te, y_tr, y_te)
```

```python
def crear_modelo(datos):
    ...
...
```

```python
grader.check("q3")
```

## Pregunta 4

```python
with open('Archivos/modeloTaller', 'rb') as f:
    modelo = pickle.load(f)
```

```python
def predecir(modelo, datos):
    ...
...
```

```python
grader.check("q4")
```

## Pregunta 5

```python
def calcular_precision(modelo, X_test, y_test):
    ...
...
```

```python
grader.check("q5")
```

## Entrega

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Asegúrate de haber ejecutado todas las celdas y que el output se muestre correctamente antes de exportar el notebook como pdf y entregarlo en Bloque Neón.

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> X, y = dividir_datos(archivo)
>>> assert isinstance(X, pd.DataFrame)
>>> assert isinstance(y, pd.Series)
```


Tests for q2 (points: 1)
```
>>> X_tr, X_te, y_tr, y_te = crear_train_test(datos)
>>> assert len(X_tr) == 4
>>> assert len(X_te) == 3
```


Tests for q3 (points: 1)
```
>>> mod = crear_modelo(datos)
>>> assert isinstance(mod, LinearRegression)
>>> assert hasattr(mod, 'coef_')
```


Tests for q4 (points: 1)
```
>>> preds = predecir(modelo, datos)
>>> assert isinstance(preds, np.ndarray)
>>> assert len(preds) == 5
```


Tests for q5 (points: 1)
```
>>> score = calcular_precision(modelo, X_te, y_te)
>>> assert np.isclose(score, -3.7894053047960226)
```
