```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller1.ipynb")
```

<img src="Archivos/miad4.png" width=800x>

# Estructuras de datos: listas, tuplas, conjuntos y diccionarios

En este taller implementarás listas, tuplas y diccionarios para resolver diversos ejercicios.


## Habilidades en práctica

Al realizar este taller podrás revisar tu progreso para:

**1.** Declarar tuplas, listas y diccionarios.<br> 
**2.** Consultar y modificar los valores almacenados en estructuras de datos y estructuras de datos anidadas.<br> 
**3.** Utilizar el rebanado en estructuras de datos ordenadas (listas y tuplas).<br> 
**4.** Utilizar métodos para modificar, reordenar, u operar estructuras de datos.<br> 


## Instrucciones

En cada uno de los siguientes ejercicios deberás escribir el código solicitado estrictamente en las celdas indicadas para ello, teniendo en cuenta las siguientes recomendaciones:

* No crear, eliminar o modificar celdas de este Notebook (salvo lo que se te indique), pues puede verse afectado el proceso de calificación automática.
* La calificación se realiza de manera automática con datos diferentes a los proporcionados en este taller. Por consiguiente, tu código debe funcionar para diferentes instancias de cada uno de los ejercicios; una instancia hace referencia al valor de los parámetros.
* La calificación de cada ejercicio depende del valor que tome la variable especificada en su enunciado. Por lo tanto, aunque declares variables adicionales en tu código, es escencial que utilices los nombres de las variables propuestas en los enunciados de los ejercicios para guardar el resultado final.

A continuación puedes ver un ejemplo de los enunciados y ejercicios que encontrarás en este taller.

## Meses del año

A continuación encuentras una lista llamada `lista_meses` que contiene todos los meses del año.

```python
# No modifiques esta celda

lista_meses = ['Enero', 'Febrero', 'Marzo', 'Abril', 'Mayo', 'Junio', 'Julio', 'Agosto', 'Septiembre', 'Octubre', 'Noviembre', 'Diciembre']
```

### Ejemplo

Extrae los últimos tres meses del año de la lista `lista_meses` y almacénalos en una nueva lista llamada `cuarto_trimestre`.

```python
# Escribe tu codigo aqui
cuarto_trimestre = lista_meses[9:]
```

## El proyecto inmobiliario 

El gerente de un proyecto inmobiliario registró en la siguiente tabla el número de casas vendidas en los últimos seis meses:

|Mes|Ventas|
|:-------:|:-----------:|
|1| 20|
|2| 21|
|3| 22|
|4| 20|
|5| 21|
|6| 27|

*Resuelve los ejercicios del 1 al 3 con base en este enunciado.*

### Ejercicio 1

Declara una lista llamada `lista_ventas` en la que se almacene la información de la columna ventas.

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q1")
```

### Ejercicio 2 

De manera inesperada, a algunos de los compradores del mes `mes_a_actualizar` le negaron el crédito. Teniendo en cuenta lo anterior, el número de casas vendidas de ese mes ha sido actualizado y está almacenado en la variable `ventas_mes_a_actualizar`. Utiliza estas variables en tu código.

```python
# No modifiques esta celda

mes_a_actualizar = 4
ventas_mes_a_actualizar = 18
```

Modifica la lista de ventas `lista_ventas` para actualizar ese valor.

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q2")
```

### Ejercicio 3

El gerente del proyecto solicitó que se le reporten los `n` mayores niveles de ventas mensuales de casas, ordenados de mayor a menor. La variable en la que guardes el resultado debe tener el nombre `mejores_ventas`. Haz uso de la lista `lista_ventas`.

**Pista:** el método `sort` recibe un parámetro llamado `reverse`.

```python
# No modifiques esta celda

n = 3
```

```python
# Escribe tu codigo aqui
...
```

## La compañía de encuestas

Una compañía de encuestas requiere hacer llamadas telefónicas a un grupo de personas y solo cuenta con dos operarios. Un operario debe llamar a la primera mitad de las personas (ordenándolas alfabéticamente), mientras que el otro llamará a la segunda mitad. La variable `nombres` almacena una lista con los nombres de las personas a llamar.

```python
# No modifiques esta celda

nombres = ["Sara", "Juan", "Alberto", "Ana", "Enrique", "Lola"]
```

```python
grader.check("q3")
```

### Ejercicio 4

Haciendo uso de la lista `nombres` y otras posibles listas auxiliares que necesites, debes crear las siguientes dos listas:

* La primera debe llamarse `primera_mitad` y contener la primera mitad de las personas en orden alfábetico. En caso de haber un número impar de personas, por ejemplo 5, en esta primera lista se incluyen las primeras tres personas.
* La segunda debe llamarse `segunda_mitad` y contener la segunda mitad de personas, según el orden alfabético.

**Pista:** para aproximar números decimales a enteros, el paquete `math` incluye las funciones `ceil` (aproxima al entero superior) y `floor` (aproxima al entero inferior).

```python
# Escribe tu codigo aqui
...
```

## Las letras del abecedario

A continuación encontrarás declarada una lista que contiene listas anidadas con las letras del alfabeto.

*Resuelve los ejercicios del 5 al 8 con base en este enunciado.*

```python
# No modifiques esta celda

lista_abecedario = ["a", "b", "c", ["d", "e", ["f", "g"]]]
```

```python
grader.check("q4")
```

### Ejercicio 5

Haciendo uso de la lista `lista_abecedario`, obtén las primeras tres letras del abecedario y almacénalas en una lista con nombre `lista_abc`. 

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q5")
```

### Ejercicio 6

Haciendo uso de la lista `lista_abecedario`, obtén una lista con las letras d y e, y almacénala en una lista con nombre `lista_de`.

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q6")
```

### Ejercicio 7

Haciendo uso de la lista `lista_abecedario` y posibles listas auxiliares que necesites, crea una lista con las letras f, g y h, y almacénala en una lista con nombre `lista_fgh`. 

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q7")
```

### Ejercicio 8

Haciendo uso de las listas anteriormente creadas y posibles listas auxiliares que necesites, genera una lista que contenga las primeras diez letras del abecedario y llámala `lista_a_j`.

```python
# Escribe tu codigo aqui
...
```

## Los números primos

Un equipo de ingenieros de sistemas está trabajando con números primos para realizar encriptaciones, en la siguiente celda encuentras los números primos del 1 al 100.

*Resuelve los ejercicios 9 y 10 con base en este enunciado.*

```python
# No modifiques esta celda

tupla_primos =  (2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47, 53, 59, 61, 67, 71, 73, 79, 83, 89, 97)
```

```python
grader.check("q8")
```

### Ejercicio 9

Encuentra cuál índice de la tupla `tupla_primos` le corresponde al número `a_buscar` y almacénalo en una variable llamada `indice`.

```python
# No modifiques esta celda

a_buscar = 73
```

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q9")
```

### Ejercicio 10

El equipo de ingenieros de sistemas decidió cambiar levemente los números a usar para la encriptación. Ahora, en vez de utilizar todos los números primos del 1 al 100, decidieron elegir un número sí, un número no, dentro de un rango definido por `lim_inferior` y `lim_superior` (incluyéndolos). Por ejemplo, si `lim_inferior` es 5 y `lim_superior` es 31, los números utilizados para la encriptación serían los siguientes: 5,11,17,23,31. 

Crea una tupla con la secuencia que necesita el equipo y llámala `tupla_secuencia`. Haz uso de la tupla `tupla_primos`.

**Pista**: a la instrucción de rebanado se le puede agregar un tercer elemento separándolo con dos puntos, para indicar a cuantos pasos rebanar la estructura de datos ordenada. Ejemplo: `tupla[indice_inicio : indice_fin : pasos]`

```python
# No modifiques esta celda

lim_inferior = 3
lim_superior = 53
```

```python
# Escribe tu codigo aqui
...
```

## La fábrica de muebles

Una fábrica de muebles hace escritorios a la medida con dos tipos de madera: caoba y arce. La empresa registró los pedidos semanales de ambas referencias de madera: en la primera semana debe fabricar tres escritorios de caoba y dos de arce, en la segunda semana necesita dos escritorios de caoba y cuatro de arce, y en la última semana un escritorio de cada tipo.

En la siguiente celda encontrarás tres tuplas que hacen referencia a los pedidos de cada semana. 

*Resuelve los ejercicios 11 y 12 con base en este enunciado.*

```python
# No modifiques esta celda

pedidos_semana_1 = (3,2)
pedidos_semana_2 = (2,4)
pedidos_semana_3 = (1,1)
```

```python
grader.check("q10")
```

### Ejercicio 11

Declara una tupla con el nombre `tupla_pedidos` que anide las tres tuplas declaradas en la anterior celda de código. La tupla `tupla_pedidos` debe quedar en órden cronológico.

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q11")
```

### Ejercicio 12

Utiliza la función de Python `sum` para obtener el total de pedidos de la semana `s` y almacena el resultado en una variable llamada `total_pedidos_semana_s`. Haz uso de la tupla `tupla_pedidos`.

```python
# No modifiques esta celda

s = 2
```

```python
# Escribe tu codigo aqui
...
```

## Meses 

En la siguiente celda encontrarás un diccionario incompleto con el último día de cada mes del año.

```python
# No modifiques esta celda

dicc_meses = {"Enero":31,
              "Febrero":28,
              "Marzo": 31,
              "Abril": 30,
              "Mayo":31,
              "Julio":31,
              "Agosto":31,
              "Septiembre":30,
              "Noviembre":30}
```

```python
grader.check("q12")
```

### Ejercicio 13

Completa el diccionario `dicc_meses` con los meses faltantes.

```python
# Escribe tu codigo aqui
...
```

## Una matriz en un diccionario 

Adaptado de Downey, A., Elkner, J., & Meyers, C. (2012).

Una tienda de videojuegos está analizando si una familia comprará la nueva consola de videojuegos con base en dos variables: el número de hijos (un hijo o más) y las horas que comparten los padres al día con sus hijos (una hora o más). La compañía representó sus resultados en la siguiente matriz, donde las columnas corresponden al número de hijos y las filas al número de horas diarias que los padres comparten con sus hijos:

$$\begin{matrix} 0 & 1 & 0 & 1 & 1 \\ 0 & 0 & 0 & 0 & 1 \\ 1 & 0 & 0 & 1 & 0 \\ 1 & 0 & 1 & 0 & 1 \end{matrix}$$

Según esta matriz, una familia en la que los padres disponen de tres horas para compartir con sus hijos y que tiene cuatro hijos comprará la consola de videojuegos. Asimismo, en una familia en la que los padres disponen de una hora para compartir con sus hijos y que tiene tres hijos no la comprará. 

*Resuelve los ejercicios 14 y 15 con base en este enunciado.*

```python
grader.check("q13")
```

### Ejercicio 14

Crea un diccionario con el nombre `matriz`, que contenga únicamente los casos en los que la familia comprará la consola. Las llaves del dicccionario deben ser tuplas que indiquen la cantidad de horas y cantidad de hijos. Además, tanto los componentes de la llave como el valor deben ser númericos. Si también representas los casos en los que la familia no comprará la consola, no obtendrás puntuación.

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q14")
```

### Ejercicio 15

Consulta el diccionario para saber si una familia, que tiene `h` hijos y en la que los padres disponen de `t` horas diarias para compartir con ellos, comprará la consola. Debes hacer uso del diccionario `matriz` y no lo puedes modificar. Guarda el valor resultante en una variable que se llame `prediccion_familia`.

**Nota**: la variable `prediccion_familia` debe tomar el valor de 1 si la familia comprará la consola, y 0 en caso contrario.

```python
# No modifiques esta celda

h = 5
t = 2
```

```python
# Escribe tu codigo aqui
...
```

## Referencias
 
Downey, A., Elkner, J., & Meyers, C. (2012). How To Think Like A Computer Scientist: Learning with Python. Segunda Edición. Recuperado de: https://www.openbookproject.net/thinkcs/python/english2e/index.html

## Créditos

__Autor(es)__: Jorge Esteban Camargo Forero, Juan Felipe Rengifo Méndez, Alejandro Mantilla Redondo, Diego Alejandro Cely Gómez
 
__Fecha última actualización__: 06/08/2021

```python
grader.check("q15")
```

## Entrega

Asegurate de haber ejecutado todas las celdas de tu notebook en orden antes de ejecutar la celda de abajo. La celda generara un archivo ZIP para que lo entregues. **Guarda tu notebook antes de exportar.**

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert type(lista_ventas) == list, f'Tu variable lista_ventas no es de tipo {list.__name__}.'
>>> assert len(lista_ventas) == 6, 'Tu lista tiene un tamafo incorrecto.'
>>> assert lista_ventas == [20, 21, 22, 20, 21, 27], 'Los valores de tu lista no son correctos.'
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q10 (points: 1)
```
>>> assert type(tupla_secuencia) == tuple
>>> assert tupla_secuencia[0] == lim_inferior, 'Tu tupla no inicia en el limite inferior.'
>>> assert tupla_secuencia[-1] == lim_superior, 'Tu tupla no termina en el limite superior.'
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q11 (points: 1)
```
>>> assert type(tupla_pedidos) == tuple
>>> assert len(tupla_pedidos) == 3
>>> for i in tupla_pedidos:
...     assert type(i) == tuple
>>> assert tupla_pedidos == (pedidos_semana_1, pedidos_semana_2, pedidos_semana_3)
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q12 (points: 1)
```
>>> assert type(total_pedidos_semana_s) == int
>>> assert total_pedidos_semana_s == sum(tupla_pedidos[s - 1])
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q13 (points: 1)
```
>>> assert type(dicc_meses) == dict
>>> assert len(dicc_meses) == 12, 'Tu diccionario no tiene los 12 meses.'
>>> for mes in ['Enero', 'Febrero', 'Marzo', 'Abril', 'Mayo', 'Junio', 'Julio', 'Agosto', 'Septiembre', 'Octubre', 'Noviembre', 'Diciembre']:
...     assert mes in dicc_meses, f'Falta el mes {mes}.'
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q14 (points: 1)
```
>>> assert type(matriz) == dict
>>> for v in matriz.values():
...     assert v == 1, 'Tu diccionario tiene casos donde la familia no compra.'
>>> assert len(matriz) == 9, 'Tu diccionario no tiene el tamano correcto.'
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q15 (points: 1)
```
>>> assert type(prediccion_familia) == int
>>> assert prediccion_familia in [0, 1]
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q2 (points: 1)
```
>>> assert type(lista_ventas) == list
>>> assert len(lista_ventas) == 6
>>> assert lista_ventas[mes_a_actualizar - 1] == ventas_mes_a_actualizar, 'No actualizaste el valor correcto.'
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q3 (points: 1)
```
>>> assert type(mejores_ventas) == list
>>> assert len(mejores_ventas) == n, 'Tu lista no tiene el tamano correcto.'
>>> assert mejores_ventas == sorted(mejores_ventas, reverse=True), 'Tu lista no esta en orden descendente.'
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q4 (points: 1)
```
>>> import math as _math
>>> assert type(primera_mitad) == list
>>> assert type(segunda_mitad) == list
>>> assert len(primera_mitad) == _math.ceil(len(nombres) / 2)
>>> assert len(segunda_mitad) == len(nombres) // 2
>>> assert primera_mitad == sorted(primera_mitad), 'primera_mitad no esta en orden alfabetico.'
>>> assert segunda_mitad == sorted(segunda_mitad), 'segunda_mitad no esta en orden alfabetico.'
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q5 (points: 1)
```
>>> assert type(lista_abc) == list
>>> assert len(lista_abc) == 3
>>> assert lista_abc == ['a', 'b', 'c']
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q6 (points: 1)
```
>>> assert type(lista_de) == list
>>> assert len(lista_de) == 2
>>> assert lista_de == ['d', 'e']
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q7 (points: 1)
```
>>> assert type(lista_fgh) == list
>>> assert len(lista_fgh) == 3
>>> assert lista_fgh == ['f', 'g', 'h']
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q8 (points: 1)
```
>>> assert type(lista_a_j) == list
>>> assert len(lista_a_j) == 10
>>> assert lista_a_j == list('abcdefghij')
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```


Tests for q9 (points: 1)
```
>>> assert type(indice) == int
>>> assert indice == tupla_primos.index(a_buscar)
>>> print('Felicidades, realizaste este ejercicio correctamente.')
```
