```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller4.ipynb")
```

# Taller 4 - Funciones, manejo de archivos y cadenas

## Pregunta 1
Sustituir caracteres.

```python
cadena = 'Archivos de texto'
```

```python
...
```

```python
grader.check("q1")
```

## Pregunta 2
Leer archivo.

```python
directorio = './Archivos/tesoro.txt'
```

```python
...
```

```python
grader.check("q2")
```

## Pregunta 3
Filtrar países.

```python
lista_paises = ['Japón', 'Francia']
lineas = archivo_a_lista('./Archivos/tesoro.txt')
```

```python
...
```

```python
grader.check("q3")
```

## Pregunta 4
ECM.

```python
pronostico_jorge = ['A', 'B']
pronostico_juan = ['B', 'A']
```

```python
...
```

```python
grader.check("q4")
```

## Entrega

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Asegurate de haber ejecutado todas las celdas...

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert sustituir('Archivos de texto') == 'Atchivos de texro'
```


Tests for q2 (points: 1)
```
>>> assert len(archivo_a_lista('./Archivos/tesoro.txt')) == 36
```


Tests for q3 (points: 1)
```
>>> assert len(filtrar_paises(['Japón', 'Francia'])) == 2
```


Tests for q4 (points: 1)
```
>>> assert ECM(pronostico_jorge, pronostico_juan) == (1.4, 2.2)
```
