```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller9.ipynb")
```

# Taller 9: Dashboard interactivo de COVID-19

En este taller construirás un dashboard interactivo usando `panel` y `plotly` con datos de COVID-19 de la OMS.

```python
# No modifiques esta celda
import panel as pn
import plotly.graph_objects as go
import pandas as pd
pn.extension('plotly')
```

## Diccionario de datos

El archivo `Archivos/COVID-19.csv` contiene las siguientes variables:
- `Name`: Nombre del país
- `WHO Region`: Región de la OMS
- `Cases - cumulative total`: Total de casos acumulados
- `Cases - cumulative total per 1 million population`: Casos acumulados por millón
- `Cases - newly reported in last 7 days`: Casos nuevos últimos 7 días
- `Cases - newly reported in last 24 hours`: Casos nuevos últimas 24 horas
- `Deaths - cumulative total`: Total de muertes acumuladas
- `Deaths - cumulative total per 1 million population`: Muertes por millón
- `Deaths - newly reported in last 7 days`: Muertes últimos 7 días
- `Deaths - newly reported in last 24 hours`: Muertes últimas 24 horas

### Paso 1: Cargar los datos

Carga el archivo `Archivos/COVID-19.csv` en un DataFrame llamado `datos`.

```python
datos = ...
...
```

```python
grader.check("q1")
```

### Paso 2: Control de tipo

Crea un widget `RadioButtonGroup` llamado `criterio_tipo` con opciones `['Casos', 'Muertes']`.

```python
criterio_tipo = ...
...
```

```python
grader.check("q2")
```

### Paso 3: Control de periodo

Crea un widget `RadioButtonGroup` llamado `criterio_periodo` con opciones `['Historico', 'Ultimos 7 dias', 'Ultimas 24 horas']`.

```python
criterio_periodo = ...
...
```

```python
grader.check("q3")
```

### Paso 4: Gráfico interactivo

Crea una función llamada `funcion_interactiva` decorada con `@pn.depends` que reciba la selección de tipo y periodo, y retorne un diagrama de barras horizontal (`go.Bar` con `orientation='h'`) mostrando los 10 países con mayor valor.

```python
# Escribe tu codigo aqui
...
```

```python
grader.check("q4")
```

### Paso 5: Muertes por millón

Crea un diagrama de barras horizontal llamado `fig_muertes_millon` con los 10 países con mayor número de muertes por millón de personas.

```python
fig_muertes_millon = ...
...
```

```python
grader.check("q5")
```

### Paso 6: Casos por región OMS

Crea un diagrama de barras horizontal llamado `fig_regiones` con los casos acumulados agrupados por región de la OMS.

```python
fig_regiones = ...
...
```

```python
grader.check("q6")
```

### Paso 7: Tablero de control

Crea el tablero de control completo llamado `tablero_control` usando `pn.Column` y/o `pn.Row` que integre los widgets, el gráfico interactivo y los gráficos estáticos.

```python
tablero_control = ...
...
```

```python
grader.check("q7")
```

## Entrega

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Asegurate de haber ejecutado todas las celdas de tu notebook antes de exportar.

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert isinstance(datos, pd.DataFrame), 'datos debe ser un DataFrame.'
>>> assert 'Name' in datos.columns, "Debe contener la columna 'Name'."
>>> assert len(datos) > 0, 'El DataFrame no debe estar vacio.'
```


Tests for q2 (points: 1)
```
>>> assert isinstance(criterio_tipo, pn.widgets.RadioButtonGroup), 'criterio_tipo debe ser RadioButtonGroup.'
>>> assert set(criterio_tipo.options) == {'Casos', 'Muertes'}, "Opciones deben ser ['Casos', 'Muertes']."
```


Tests for q3 (points: 1)
```
>>> assert isinstance(criterio_periodo, pn.widgets.RadioButtonGroup), 'criterio_periodo debe ser RadioButtonGroup.'
>>> assert len(criterio_periodo.options) == 3, 'Debe tener 3 opciones.'
```


Tests for q4 (points: 1)
```
>>> assert callable(funcion_interactiva), 'funcion_interactiva debe ser callable.'
>>> fig_test = funcion_interactiva('Casos', 'Historico')
>>> assert isinstance(fig_test, go.Figure), 'Debe retornar un go.Figure.'
```


Tests for q5 (points: 1)
```
>>> assert isinstance(fig_muertes_millon, go.Figure), 'fig_muertes_millon debe ser go.Figure.'
>>> assert fig_muertes_millon.data[0].orientation == 'h', 'Debe ser horizontal.'
```


Tests for q6 (points: 1)
```
>>> assert isinstance(fig_regiones, go.Figure), 'fig_regiones debe ser go.Figure.'
>>> assert fig_regiones.data[0].orientation == 'h', 'Debe ser horizontal.'
```


Tests for q7 (points: 1)
```
>>> assert tablero_control is not None, 'tablero_control no debe ser None.'
```
