```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller6.ipynb")
```

# Taller 6: Explorar y modificar bases de datos utilizando `pandas`

En este taller filtrarás y modificarás bases de datos por medio de sus índices al resolver diversos ejercicios.

## Instrucciones

En cada uno de los siguientes ejercicios deberás implementar la función solicitada. La calificación se realiza de manera automática.

```python
# No modifiques esta celda
import numpy as np
import pandas as pd
import inspect
```

### Ejercicio 1

Junto a este Notebook encontrarás un archivo llamado `"pooled_data_public.dta"` que contiene los resultados de la encuesta de hogares.

Implementa una función llamada `importar_e_indexar` que no reciba parámetros, importe los datos del archivo `"Archivos/pooled_data_public.dta"` en un `DataFrame`, establezca la columna `"id"` como el índice del `DataFrame` y lo retorne.

*Esta función debe retornar un `DataFrame` de 230509 filas y 128 columnas.*

```python
def importar_e_indexar():
    # Escribe tu codigo aqui
    ...
...
```

```python
grader.check("q1")
```

### Ejercicio 2

Implementa una función llamada `indice_sencillo_a_multiple` que no reciba parámetros, tome el `DataFrame` resultante de la función `importar_e_indexar` y cambie el índice de uno sencillo a uno múltiple al agregar la columna `"country_code"` al índice existente.

*Esta función debe retornar un `DataFrame` con un `MultiIndex` de 230509 filas y 127 columnas.*

```python
def indice_sencillo_a_multiple():
    # Escribe tu codigo aqui
    ...
...
```

```python
grader.check("q2")
```

### Ejercicio 3

Implementa una función llamada `sin_columnas` que no reciba parámetros, tome el `DataFrame` resultante de la función `importar_e_indexar` y elimine las columnas `"country_name"`, `"grado"` y `"coima"` del `DataFrame`. Utiliza el método `.drop()`.

*Esta función debe retornar un `DataFrame` de 230509 filas y 125 columnas.*

```python
def sin_columnas():
    # Escribe tu codigo aqui
    ...
...
```

```python
grader.check("q3")
```

### Ejercicio 4

Implementa una función llamada `indexar_columnas_numericamente` que no reciba parámetros, tome el `DataFrame` resultante de la función `importar_e_indexar` y elimine las columnas que se encuentran en las posiciones 80 a 100 (incluyendo la posición 100).

*Esta función debe retornar un `DataFrame` de 230509 filas y 107 columnas.*

```python
def indexar_columnas_numericamente():
    # Escribe tu codigo aqui
    ...
...
```

```python
grader.check("q4")
```

### Ejercicio 5

A raíz de la pandemia, los gobiernos de latinoamérica desean monitorear a aquellas personas que trabajan fuera de su casa y que tienen adultos mayores en su hogar.

Implementa una función llamada `trabajo_afuera_y_mayores` que no reciba parámetros, tome el `DataFrame` resultante de la función `importar_e_indexar` y filtre aquellas personas que trabajan fuera de su casa (`trabajocasa == "Fuera de su casa"`) y que tienen adultos mayores en su hogar (`haymayores == "si"`).

*Esta función debe retornar un `DataFrame` de 14311 filas.*

```python
def trabajo_afuera_y_mayores():
    # Escribe tu codigo aqui
    ...
...
```

```python
grader.check("q5")
```

### Ejercicio 6

Implementa una función llamada `hogares_decrecen` que no reciba parámetros, tome el `DataFrame` resultante de la función `importar_e_indexar` y filtre aquellos hogares donde el número de personas en el hogar decreció durante la pandemia (`nrhogar_antes > nrhogar`).

*Esta función debe retornar un `DataFrame`.*

```python
def hogares_decrecen():
    # Escribe tu codigo aqui
    ...
...
```

```python
grader.check("q6")
```

## Entrega

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Asegurate de haber ejecutado todas las celdas de tu notebook antes de exportar.

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> res = importar_e_indexar()
>>> assert isinstance(res, pd.DataFrame), 'Debe retornar un DataFrame.'
>>> assert res.shape == (230509, 128), f'Dimensiones incorrectas: {res.shape}, esperado (230509, 128).'
>>> assert res.index.name == 'id', "El indice debe ser la columna 'id'."
```


Tests for q2 (points: 1)
```
>>> res2 = indice_sencillo_a_multiple()
>>> assert isinstance(res2, pd.DataFrame), 'Debe retornar un DataFrame.'
>>> assert isinstance(res2.index, pd.MultiIndex), 'El indice debe ser un MultiIndex.'
>>> assert res2.shape == (230509, 127), f'Dimensiones incorrectas: {res2.shape}.'
```


Tests for q3 (points: 1)
```
>>> res3 = sin_columnas()
>>> assert isinstance(res3, pd.DataFrame), 'Debe retornar un DataFrame.'
>>> assert res3.shape == (230509, 125), f'Dimensiones incorrectas: {res3.shape}.'
>>> assert 'country_name' not in res3.columns, "La columna 'country_name' no debe existir."
```


Tests for q4 (points: 1)
```
>>> res4 = indexar_columnas_numericamente()
>>> assert isinstance(res4, pd.DataFrame), 'Debe retornar un DataFrame.'
>>> assert res4.shape == (230509, 107), f'Dimensiones incorrectas: {res4.shape}, esperado (230509, 107).'
```


Tests for q5 (points: 1)
```
>>> res5 = trabajo_afuera_y_mayores()
>>> assert isinstance(res5, pd.DataFrame), 'Debe retornar un DataFrame.'
>>> assert res5.shape[0] == 14311, f'Cantidad de filas incorrecta: {res5.shape[0]}, esperado 14311.'
```


Tests for q6 (points: 1)
```
>>> res6 = hogares_decrecen()
>>> assert isinstance(res6, pd.DataFrame), 'Debe retornar un DataFrame.'
>>> assert res6.shape[0] > 0, 'El DataFrame no debe estar vacio.'
```
