```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller8.ipynb")
```

# Taller 8: Visualización e Interactividad

En este taller crearás gráficos interactivos usando `plotly` e `ipywidgets`.

```python
# No modifiques esta celda
import plotly.graph_objects as go
import ipywidgets as widgets
import pandas as pd
from ipywidgets import interactive
import numpy as np
```

### Ejercicio 1

Utiliza `plotly` para realizar un diagrama de torta (`go.Pie`) sobre el estado conyugal (`Marital_Status`) de los clientes de la base de datos `BankChurners.csv`. Resalta la categoría más predominante usando el parámetro `pull`.

Guarda tu gráfico en una variable llamada `fig`.

```python
# No modifiques esta celda
df = pd.read_csv('Archivos/BankChurners.csv')
```

```python
# Escribe tu codigo aqui
fig = ...
...
```

```python
grader.check("q1")
```

### Ejercicio 2

Utiliza `plotly` para realizar un histograma (`go.Histogram`) que permita verificar si el cambio en el monto de transacciones (`Total_Amt_Chng_Q4_Q1`) se ajusta a una distribución normal.

Guarda tu gráfico en una variable llamada `fig`.

```python
# Escribe tu codigo aqui
fig = ...
...
```

```python
grader.check("q2")
```

### Ejercicio 3

Utiliza `plotly` para realizar un gráfico de dispersión (`go.Scatter` con `mode='markers'`) que permita visualizar la relación entre las variables `Year_Birth` y `MntWines` de la base de datos de mercadeo.

Guarda tu gráfico en una variable llamada `fig`.

```python
# No modifiques esta celda
mercadeo = pd.read_csv('Archivos/marketing_data.csv')
```

```python
# Escribe tu codigo aqui
fig = ...
...
```

```python
grader.check("q3")
```

### Ejercicio 4

Utiliza `ipywidgets` para filtrar las observaciones de la base de datos de mercadeo de acuerdo con los siguientes criterios:
1. Nivel máximo de escolaridad alcanzado (`Education`).
2. Número de niños en el hogar (`Kidhome`).
3. Número de adolescentes en el hogar (`Teenhome`).

Crea una función llamada `filtrar` y un widget interactivo llamado `interactive_plot`.

```python
# Escribe tu codigo aqui
def filtrar(Escolaridad, Ninos, Adolescentes):
    ...

interactive_plot = ...
...
```

```python
grader.check("q4")
```

### Ejercicio 5

Utiliza `ipywidgets` y `plotly` para generar un gráfico con controles que:
1. Cuente con un control desplegable para seleccionar una de las variables de gastos frente a los ingresos.
2. Genere un gráfico de dispersión de la variable seleccionada vs `Income`.
3. Genere un diagrama de caja de la variable seleccionada.

Variables de gastos: `MntWines`, `MntFruits`, `MntMeatProducts`, `MntFishProducts`, `MntSweetProducts`, `MntGoldProds`.

Crea una función llamada `generar_graficos` y un widget llamado `control_gastos`.

```python
# Escribe tu codigo aqui
def generar_graficos(variable_gasto='MntWines'):
    ...

control_gastos = ...
...
```

```python
grader.check("q5")
```

## Entrega

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Asegurate de haber ejecutado todas las celdas de tu notebook antes de exportar.

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert isinstance(fig, go.Figure), 'fig debe ser un go.Figure.'
>>> assert len(fig.data) > 0, 'La figura debe contener al menos un trazo.'
>>> assert isinstance(fig.data[0], go.Pie), 'El trazo debe ser un go.Pie.'
```


Tests for q2 (points: 1)
```
>>> assert isinstance(fig, go.Figure), 'fig debe ser un go.Figure.'
>>> assert isinstance(fig.data[0], go.Histogram), 'El trazo debe ser un go.Histogram.'
```


Tests for q3 (points: 1)
```
>>> assert isinstance(fig, go.Figure), 'fig debe ser un go.Figure.'
>>> assert isinstance(fig.data[0], go.Scatter), 'El trazo debe ser un go.Scatter.'
>>> assert 'markers' in (fig.data[0].mode or 'markers'), 'El modo debe incluir markers.'
```


Tests for q4 (points: 1)
```
>>> assert callable(filtrar), 'filtrar debe ser una funcion.'
>>> assert isinstance(interactive_plot, widgets.Widget), 'interactive_plot debe ser un widget.'
```


Tests for q5 (points: 1)
```
>>> assert callable(generar_graficos), 'generar_graficos debe ser una funcion.'
>>> assert isinstance(control_gastos, widgets.Widget), 'control_gastos debe ser un widget.'
```
