```python
# Initialize Otter
import otter
grader = otter.Notebook("Taller5.ipynb")
```

```python
import numpy as np
import pandas as pd
```

## Ejercicio 1

```python
mercadeo = np.array([2000, 1500, 1300, 1600])
logistica = np.array([1800, 1600, 1200, 1600])
```

```python
...
```

```python
grader.check("q1")
```

## Ejercicio 2

```python
hermanos = pd.Series({'Juan': 45, 'Maria': 42, 'Pedro': 47, 'Ana': 39})
```

```python
...
```

```python
grader.check("q2")
```

## Contexto para Ejercicios 3, 4 y 5

```python
empleados_data = {'Ciudad Residencia': ['Bogota', 'Medellin', 'Cali', 'Bogota', 'Barranquilla'], 'Salario (en USD)': [1200, 1100, 1050, 1500, 1300]}
empleados = pd.DataFrame(empleados_data, index=['Juan', 'Maria', 'Pedro', 'Ana', 'Luis'])
```

```python
...
```

```python
grader.check("q3")
```

```python
...
```

```python
grader.check("q4")
```

```python
...
```

```python
grader.check("q5")
```

## Contexto para Ejercicios 6, 7 y 8

```python
paises = pd.read_csv('./Archivos/DiccionarioPaises.csv', encoding='UTF-8', delimiter=';')
alfabeto = ('A','B','C','D','E','F','G','H','I','J','K','L','M','N','Ñ','O','P','Q','R','S','T','U','V','W','X','Y','Z')
```

```python
...
```

```python
grader.check("q6")
```

```python
...
```

```python
grader.check("q7")
```

```python
...
```

```python
grader.check("q8")
```

## Entrega

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Asegurate de haber ejecutado todas las celdas...

```python
# Guarda tu notebook y ejecuta esta celda para exportar tu entrega.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert np.isclose(promedio(mercadeo, logistica), 1575.0)
```


Tests for q2 (points: 1)
```
>>> assert mayores(hermanos) == (47, 45)
```


Tests for q3 (points: 1)
```
>>> assert ciudades(empleados) == 4
```


Tests for q4 (points: 1)
```
>>> assert mayor_al_promedio(empleados, 'Juan') is False
```


Tests for q5 (points: 1)
```
>>> assert 'Sueldo (en USD)' in modificar(empleados).columns
```


Tests for q6 (points: 1)
```
>>> assert dif_longitud(paises).shape[0] == 127
```


Tests for q7 (points: 1)
```
>>> assert paises_por_a(paises).shape[0] == 13
```


Tests for q8 (points: 1)
```
>>> assert inicia_con(paises)['B'] == 18
```
